<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 17054722  -->
  <question type="stack">
    <name>
      <text>Absolute Value In Lpp</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[Reformulate the problem
\[
\begin{array}{lrcrcr}
\mathrm{minimize} & 2x_1 &+& 3|x_2 &-& {@a@}| \\
\mathrm{subject\ to\ } & |x_1 + 2| &+& x_2 &\leq& 5,
\end{array}
\]
as a linear programming problem by replacing the argument \(x_k\) of each absolute value \(|x_k|\) as the difference of two new non-negative decision variables \(p_k\) and \(m_k\), expressing its absolute value as their sum. Then
put the problem into standard form, by introducing \(s\) slack variables \(x_k\) as needed, where \(k=n+1\ldots n+s\) and \(n\) is the number of original decision variables. Please retain all constants in the cost function, so that the standard-form cost agrees with the cost in the original problem.
<p>
\(\mathrm{minimize}\) [[input:cost]] \(\mathrm{subject\ to}\)
<p>
[[input:constraint]],
<p>
[[input:variables]]\(\ge\mathbf{0}\),
<p>
[[validation:cost]]
<p>
[[validation:constraint]]
<p>
[[validation:variables]]]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[Let \(x_1+2=p_1-m_1\) and replace \(|x_1+2|\) by \(p_1+m_1\).
Let \(x_2-{@a@}=p_2-m_2\), replacing \(|x_2-{@a@}|\) with
\(p_2+m_2\). We obtain the equivalent linear programming problem
\[
\begin{array}{lrcrcrcrcrcr}
\mathrm{minimize}	&2p_1	&-&2m_1	&+&3p_2	&+&3m_2	&-&4  \\
\mathrm{subject\ to} 	&p_1	&+&m_1	&+&p_2	&-&m_2  &+&x_3	&=&{@5-a@},\\
			&[p_1,\kern-2.78pt&&m_1,\kern-2.78pt&&p_2,\kern-2.78pt&&m_2,\kern-2.78pt&&x_3]&\ge& \mathbf{0}.
\end{array}
\]]]></text>
    </generalfeedback>
    <defaultgrade>3</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2019121800</text>
    </stackversion>
    <questionvariables>
      <text>ordergreat(p1,m1,p2,m2,x3)
a:rand_with_step(1,20,1)
_cost:2*p1-2*m1+3*p2+3*m2-4
_constraint:p1+m1+p2-m2+x3=5-a
_variables:[p1,m1,p2,m2,x3]</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]
[[feedback:constraint]]
[[feedback:variables]]</text>
    </specificfeedback>
    <questionnote>
      <text>Variables: a
Any edits should be made to the original source: C:\Users\zhouc\Python_Workspace\tools\examples\absolute-value-in-lpp.py
This question was generated with qcreate version 0.98 from https://gitlab.com/stacktools/tools</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>1</assumepositive>
    <assumereal>1</assumereal>
    <prtcorrect format="html">
      <text>Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Incorrect answer.</text>
    </prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>constraint</name>
      <type>algebraic</type>
      <tans>_constraint</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>[[BASIC-ALGEBRA]],[[BASIC-CALCULUS]],[[BASIC-MATRIX]]</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>cost</name>
      <type>algebraic</type>
      <tans>_cost</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>[[BASIC-ALGEBRA]],[[BASIC-CALCULUS]],[[BASIC-MATRIX]]</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>variables</name>
      <type>algebraic</type>
      <tans>_variables</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>[[BASIC-ALGEBRA]],[[BASIC-CALCULUS]],[[BASIC-MATRIX]]</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>constraint</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>constraint</sans>
        <tans>_constraint</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>constraint-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>constraint: correct</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>constraint-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>constraint: incorrect</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>simplify(cost-_cost)</sans>
        <tans>0</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>cost: correct</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>cost: incorrect</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>variables</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>variables</sans>
        <tans>_variables</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>variables-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>variables: correct</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>variables-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>variables: incorrect</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <hint format="html">
      <text><![CDATA[<p>Please try again. You have 4 attempts remaining.<br></p>]]></text>
    </hint>
    <hint format="html">
      <text><![CDATA[<p>Please try again. You have 3 attempts remaining.<br></p>]]></text>
    </hint>
    <hint format="html">
      <text><![CDATA[<p>Please try again. You have 2 attempts remaining.<br></p>]]></text>
    </hint>
    <hint format="html">
      <text><![CDATA[<p>Please try again. You have 1 attempt remaining.<br></p>]]></text>
    </hint>
    <tags>
      <tag><text>tag1</text>
</tag>
      <tag><text>tag2</text>
</tag>
    </tags>
  </question>

</quiz>